```python
# Initialize Otter
import otter
grader = otter.Notebook("discussion.ipynb")
```

# Discussion 7

### Due Saturday November 6th, 11:59:59PM


# Web Scraping and Regex

```python
import os
import numpy as np
import pandas as pd
import requests
import time
from bs4 import BeautifulSoup
import re
```

```python
from discussion import *
```

## Let's scrape some books data!
* We will be scraping some book information from http://books.toscrape.com/index.html
* Note that this website is built for scraping purposes

### Get all the book urls from single catalogue page

```python
url = 'http://books.toscrape.com/catalogue/page-1.html'
```

```python
book_request = requests.get(url)
soup = BeautifulSoup(book_request.text)
soup
```

```python
articles = soup.find_all('article', {"class": "product_pod"})
articles
```

```python
book_links = [art.find('a').get('href') for art in articles]
book_links
```

### Parse Book information from a book URL

```python
# For the first book
book_url = 'http://books.toscrape.com/catalogue/' + book_links[0]
```

```python
book_soup = BeautifulSoup(requests.get(book_url).text)
book_soup
```

```python
name = book_soup.find('article', {'class':'product_page'}).h1.text
name
```

```python
price = book_soup.find('p', {'class':'price_color'}).text
price
```

```python
price = float(price[2:])
price
```

```python
rating = book_soup.find('p', class_ = re.compile('star-rating')).get('class')[1]
rating
```

**Question 1**

Get all the catalog urls of the website in the form of a list. Notice that each page of the website has a very specific pattern.

```python

```

```python

```

```python
grader.check("q1")
```

**Question 2**

Get all the category names (`Travel`, `Mystery`..) as a list. Notice to remove all the extra whitespaces if any.

```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
cats = book_categories()
```

```python
grader.check("q2")
```

# Regular Expressions

## Useful links

https://pythex.org/

https://www.debuggex.com/cheatsheet/regex/python

https://www.dataquest.io/wp-content/uploads/2019/03/python-regular-expressions-cheat-sheet.pdf

**Question 3**: Identify duplicate words in a sentence

Given an input sentence, provide a list of words that are duplicated. If there is no duplication return an empty list

```python
# Use capture group number to identify duplicates
```

```python

```

```python
grader.check("q3")
```

**Question 4**: Extract laptop specifications
The below image provides details on column names and the extract patterns. Given a df with product description - Return df with added columns of `processor` (i3, i5), `generation` (9th Gen, 10th Gen), `storage` (512 GB SSD, 1 TB HDD), `display_in_inch` (15.6 inch, 14 inch). 

If there is no specific information present, keep a null (`NaN`) value.

<img src='imgs/laptop_specs.PNG'>

```python
# Use 'pd_column.str.extract(r'pattern')' to extract the required pattern
df = pd.read_csv('data/laptop_details.csv')
```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
out = laptop_details(pd.read_csv('data/laptop_details.csv'))
```

```python
grader.check("q4")
```

## Congratulations! You're done!

* Submit your `.py` file to Gradescope. Note that you only need to submit the `.py` file; this notebook should not be uploaded. Make sure that all of your work is in the `.py` file and not here by running the doctests: `python -m doctest discussion.py`.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> urls = get_website_urls()
>>> len(urls)==50
True
```
failure_message: doctest
```
>>> urls = get_website_urls()
>>> page_nums = [int(re.findall('page-(\d)',url)[0]) for url in urls]
>>> sum(page_nums)==150
True
```
failure_message: sum of all page numbers


Tests for q2 (points: None)
```
>>> len(cats)==50
True
```
failure_message: doctest
```
>>> 'Classics' in cats
True
```
failure_message: doctest
```
>>> min(cats)=='Academic' and max(cats)=='Young Adult'
True
```
failure_message: max and min by alphabet sort


Tests for q3 (points: None)
```
>>> res = duplicate_words('I went to the market market with my my my family')
>>> res == ['market', 'my']
True
```
failure_message: doctest: multiple duplicates
```
>>> res = duplicate_words('necessity is the mother of invention')
>>> res == []
True
```
failure_message: doctest: no duplicates
```
>>> res = duplicate_words('fair fairer fairest')
>>> res == []
True
```
failure_message: doctest: word boundary


Tests for q4 (points: None)
```
>>> out.shape == (21, 5)
True
```
failure_message: doctest: shape
```
>>> out['generation'].nunique() == 4
True
```
failure_message: doctest: unique generations
```
>>> list(out.iloc[0].values)[1:] == ['i5', '9th Gen', '512 GB SSD', '15.65 inch']
True
```
failure_message: doctest: first row
```
>>> avg_display_inch = out['display_inch'].str.split().str[0].astype(float).mean()
>>> np.isclose(avg_display_inch, 14.795000000000002, atol=0.01)
True
```
failure_message: doctest: average display inches
